# Denison DA210/CS181 Homework 3.a - Step 1

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import json
import os
import os.path
import io
import sys
from contextlib import redirect_stdout

datadir = "publicdata"
```

---

## Part A: Basic JSON file processing

**Q1:** In the data directory is a file named `fib.json` containing a JSON-encoded list of numbers from the Fibonacci sequence.  Write a function
```
    readJsonList(filepath)
```
that reads and returns a list from `filepath`, **assuming a JSON-encoded text file**.  If `filepath` does not exist, or if the result is a data structure other than a list, return the empty list. (Hint: the Python function `isinstance()` can be helpful to verify the type of a data structure, and you may find the function [`os.path.isfile()`](https://docs.python.org/3/library/os.path.html#os.path.isfile) handy.)

Then use your function to read in the data from `fib.json` in `datadir` and assign the result to `json_list_1`.

<!-- grade_id: cell-d2bae836fbc119c2, points: 0, grade: false, solution: true -->
```python
def readJsonList(filepath):
    """
    Opens and reads the input filepath and returns a JSON-econded list in the filepath. Returns an empty list with filepath does not exist or the file inside filepath is not a list. 
    """ 
    if not os.path.isfile(filepath): 
        return []
    f = open(filepath, "r")
    hi = json.load(f)
    if not isinstance(hi, list):
        return []
    f.close()
    return hi

filepath = os.path.join(datadir, "fib.json")
json_list_1 = readJsonList(filepath)
json_list_1
```

Output:
```
[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89]
```

<!-- grade_id: cell-82c0e42c81213cef, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert readJsonList.__doc__ is not None # don't forget the docstring

assert json_list_1[-1] == 89
assert len(json_list_1) == 12

json_list_2 = readJsonList("./foobar.json") # does not exist
assert json_list_2 == []

json_list_3 = readJsonList(os.path.join(datadir, "config.json")) # is not a list
assert json_list_3 == []
```

**Q2:** Write a function
```
    writeSquares(n, filepath)
```
that creates a list of the squares of integers from `1` to `n`, inclusive, and then writes them out as a JSON-encoded text file to the file location given by `filepath`.

If `n` is less than or equal to `0`, no file should be created.  This function returns no value.

<!-- grade_id: cell-469c714260019773, points: 0, grade: false, solution: true -->
```python
def writeSquares(n, filepath): 
    """
    Takes in a number n and a filepath. Creates a list of all squared numbers from 1 to n, for n>0. Saves that list as a JSON list inside the input filepath.
    """
    if n> 0: 
        square = []
        for x in range(1, n+1):
            square.append(x**2)
        f = open(filepath, "w") 
        hi = json.dump(square, f)
        f.close()
        return hi
    

```

<!-- grade_id: cell-e64a97ee03c45a63, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert writeSquares.__doc__ is not None # don't forget the docstring

filepath = "./squares.json"
writeSquares(5, filepath)
result = json.load(open(filepath, 'r'))
assert result == [1, 4, 9, 16, 25]
writeSquares(1, filepath)
result = json.load(open(filepath, 'r'))
assert result == [1]
os.remove(filepath)
writeSquares(0, filepath)
assert not os.path.isfile(filepath)
```

---

## Part B: JSON data types

**Q3:** In the data directory is a file named `baby_2015_female_namecount.txt` with a line for each of the top female baby names of 2015.  Each line contains tab-separated values for the name and the count of US Social Security applications for that name.

Write a function
```
    readNameCount(filepath)
```
that reads a file formatted this way (note that this is not a JSON format), given by `filepath` and returns a tuple of two lists, the first being the names and the second being the counts (stored as integers).  The function should behave properly regardless of how many name-count lines are in the file.

If the specified `filepath` does not exist, your function should return `None`.

<!-- grade_id: cell-00f558d6134a208b, points: 0, grade: false, solution: true -->
```python
def readNameCount(filepath): 
    """
    Reads the input filepath and creates a list of names and a list of counts, then returns a tuple of the two lists.
    """
    try: 
        f = open(filepath, "r")
        name = []
        count =[]
        for lines in f: 
            line = lines.split("	")
            name.append(line[0])
            count.append(int(line[1]))
        f.close()
        return (name, count)
    except:
        return None

```

<!-- grade_id: cell-8d778e7c955efac5, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert readNameCount.__doc__ is not None # don't forget the docstring

assert readNameCount("./foo") == None
namecount_result = readNameCount(os.path.join(datadir, "baby_2015_female_namecount.txt"))
assert len(namecount_result) == 2
assert len(namecount_result[0]) == 10
assert len(namecount_result[1]) == 10
```

**Q4:** Write code that *uses* your `readNameCount` function to obtain the lists and **writes the result** to a **JSON-encoded** text file in the current directory named `namecount.json`.

<!-- grade_id: cell-00f558d6134a208b, points: 0, grade: false, solution: true -->
```python

namecount_result = readNameCount(os.path.join(datadir, "baby_2015_female_namecount.txt"))

file = open("namecount.json", "w") 
json.dump(namecount_result, file)
file.close()


```

<!-- grade_id: cell-8d778e7c955efac5, points: 3, grade: true, solution: false -->
```python
# Testing cell
jsonpath = os.path.join(".", "namecount.json")
assert os.path.isfile(jsonpath)
with open(jsonpath, 'r') as f:
    namecount_result2 = json.load(f)
assert len(namecount_result2) == 2
assert len(namecount_result2[0]) == 10
assert len(namecount_result2[1]) == 10
```

**Q5:** Consider the table of data:
```
    subject | name               | department
    --------|--------------------|--------------
    CS      | Computer Science   | CS
    MATH    | Mathematics        | MATH
    ENGL    | English Literature | ENGL
```

This could be represented in a Python program as a *dictionary* mapping from a *key*, `"subjects"`, to a *list* of dictionaries, one per subject.  These inner dictionaries map from `"subject"`, `"name"`, and `"department"` to the respective value for the row.  Be careful and follow the above specification **exactly**.

Write Python code to construct the above representation as a Python in-memory data structure called `ds`, and then have your code **encode the data structure** into a JSON-formatted string called `s`.

<!-- grade_id: cell-0abd378083a7bdd9, points: 0, grade: false, solution: true -->
```python
ds = {"subjects": [{"subject":"CS", "name": "Computer Science","department":"CS"},
      {"subject":"MATH", "name":"Mathematics", "department":"MATH"}, 
      {"subject":"ENGL", "name":"English Literature", "department":"ENGL"}]}

s = json.dumps(ds)
# Display the string
print(s)
```

Output:
```
{"subjects": [{"subject": "CS", "name": "Computer Science", "department": "CS"}, {"subject": "MATH", "name": "Mathematics", "department": "MATH"}, {"subject": "ENGL", "name": "English Literature", "department": "ENGL"}]}
```

<!-- grade_id: cell-765331274d04ea6d, points: 3, grade: true, solution: false -->
```python
# Testing Cell
f = io.StringIO()
with redirect_stdout(f):
    print(s)
    
assert f.getvalue() == '{"subjects": [{"subject": "CS", "name": "Computer Science", "department": "CS"}, {"subject": "MATH", "name": "Mathematics", "department": "MATH"}, {"subject": "ENGL", "name": "English Literature", "department": "ENGL"}]}\n'
```

**Q6:** Again encode the subjects dictionary as a JSON-formatted string, but include `indent=2` as an argument in the conversion, and assign to `s2`.

```python
s2 = json.dumps(ds, indent = 2)

# Display the string
print(s2)
```

Output:
```
{
  "subjects": [
    {
      "subject": "CS",
      "name": "Computer Science",
      "department": "CS"
    },
    {
      "subject": "MATH",
      "name": "Mathematics",
      "department": "MATH"
    },
    {
      "subject": "ENGL",
      "name": "English Literature",
      "department": "ENGL"
    }
  ]
}
```

<!-- grade_id: cell-765331274d04ea6d, points: 3, grade: true, solution: false -->
```python
# Testing Cell
f2 = io.StringIO()
with redirect_stdout(f2):
    print(s2)

assert f2.getvalue() == '{\n  "subjects": [\n    {\n      "subject": "CS",\n      "name": "Computer Science",\n      "department": "CS"\n    },\n    {\n      "subject": "MATH",\n      "name": "Mathematics",\n      "department": "MATH"\n    },\n    {\n      "subject": "ENGL",\n      "name": "English Literature",\n      "department": "ENGL"\n    }\n  ]\n}\n'
```

---

---

## Part C

**Q7:** How much time (in minutes/hours) did you spend on this homework assignment

40 minutes

**Q8:** Who was your partner for this assignment?  If you worked alone, say so instead.

Alone