# Denison DA210/CS181 Homework 2.g - Step 2

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import pandas as pd

datadir = "publicdata"
```

---

## Part D: More data normalization

**Q8:** Read `us_rent_income.csv` into a dataframe (with no specified index) named `df_rent_untidy`.

<!-- grade_id: cell-5bba2b9d91e4bf79, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display part of the untidy DataFrame
df_rent_untidy.head()
```

<!-- grade_id: cell-d927d8a1051f9d9d, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert df_rent_untidy.shape == (104,5)
assert df_rent_untidy.iloc[0,2] == "income"
assert df_rent_untidy.iloc[1,3] == 747.0
```

**Q9:** Transform `df_rent_untidy` as needed to make it tidy. Store the result as `df_rent`.  Make sure to keep the original `df_rent_untidy` unchanged.

Hint: Your tidy result should have only one column used as the row index.

<!-- grade_id: cell-5bba2b9d91e4bf79, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display part of the tidy DataFrame
df_rent.head()
```

<!-- grade_id: cell-d927d8a1051f9d9d, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert df_rent.shape == (52,5)
assert df_rent.iloc[0,1] == 24476.0
assert df_rent.iloc[0,2] == 747.0
assert df_rent.iloc[0,3] == 136.0
assert df_rent.iloc[0,4] == 3.0
assert df_rent.iloc[20,1] == 37147.0
assert df_rent.iloc[31,2] == 809.0
```

---

---

## Part E

**Q10:** How much time (in minutes/hours) did you spend on this homework assignment?

YOUR ANSWER HERE

**Q11:** Who was your partner for this assignment?  If you worked alone, say so instead.

YOUR ANSWER HERE