# Denison DA210/CS181 Homework 1.f - Step 2

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path

datadir = "publicdata"
```

---

## Part B

**Q4:** Assume that `path` refers to a CSV file with the same format used in our variations of `topnames` with a header line containing the four column names, and then rows of comma-separated field values.  Write a function

```
    readTopNamesLoD(path)
```

that reads the file and creates a LoD representation and returns the structure.

<!-- grade_id: cell-3ac611d5fbae072d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Debugging cell
tn10data = readTopNamesLoD(os.path.join(datadir, "tn10.csv"))
print(tn10data)
```

<!-- grade_id: cell-ab5935ca57c293ca, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert readTopNamesLoD.__doc__ is not None # don't forget the docstring!

tn10data = readTopNamesLoD(os.path.join(datadir, "tn10.csv"))

assert isinstance(tn10data, list)
assert len(tn10data) == 10

assert len(tn10data[0]) == 4
assert isinstance(tn10data[0], dict)

assert tn10data[0]['sex'] == 'Female'
assert tn10data[0]['year'] == 2014
assert tn10data[0]['name'] == 'Emma'
assert tn10data[0]['count'] == 20936
```

**Q5:** Write a function 

```
    filterTopNamesLoD(tnLoD, threshold)
```

to create a **filtered** copy of a topnames LoD `tnLoD` (with columns `year`, `sex`, `name`, `count`) so that only rows with a `count` value greater than or equal to `threshold` are present in the newly created LoD.

Your function should return the new and filtered LoD, nd not modify the original `tnLoD`.

```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Debugging cell
tn10 = readTopNamesLoD(os.path.join(datadir, "tn10.csv"))
tn10_filter = filterTopNamesLoD(tn10, 19000)
print(tn10_filter)
```

```python
# Testing cell
assert filterTopNamesLoD.__doc__ is not None # don't forget the docstring!

LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]

filtered = filterTopNamesLoD(LoD, 19000)
assert len(filtered) == 8
assert len(filtered[0]) == 4

assert len(LoD) == 10
assert len(LoD[0]) == 4
```

**Q6:** Write a function 

```
    addCatColumnLoD(tnLoD, threshold1, threshold2)
```

that adds a categorical column to a LoD representation in parameter `tnLoD` with the new column named `category` whose values are the strings `"small"` when `count` is below `threshold1`, is `"medium"` when `count` is greater than or equal to `threshold1` and less than `threshold2`, and `"large"` when `count` is greater than or equal to `threshold2`.

This change to tnLoD happens in place, rather than creating a new list, and so nothing is returned from the function.

<!-- grade_id: cell-0d9878d3fcbb29fa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Debugging cell
LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]

addCatColumnLoD(LoD, 19000, 19500)
print(LoD)
```

<!-- grade_id: cell-1da7403c4d9cd8aa, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert addCatColumnLoD.__doc__ is not None # don't forget the docstring!

LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]

res = addCatColumnLoD(LoD, 19000, 19500)
assert res is None

assert len(LoD) == 10
assert isinstance(LoD[0], dict)
assert len(LoD[0]) == 5
assert 'category' in LoD[0]
assert [ rowD['category'] for rowD in LoD ] == ['large', 'medium', 'large',
 'large', 'medium', 'medium', 'large', 'small', 'small', 'large']
```

---

---
## Part C

How much time (in minutes/hours) did you spend on this homework assignment?

YOUR ANSWER HERE

Who was your partner for this assignment?  If you worked alone, say so instead.

YOUR ANSWER HERE