# Denison DA210/CS181 Homework 1.e - Step 1

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import io
import sys
import string
import math
from contextlib import redirect_stdout
```

---

## Part A: Sorting

The following set of questions involve the `sorted()` built-in Python function and the list of tuples:

```python
students = [('bob',   'B', 19),
            ('chris', 'B', 18),
            ('chris', 'A', 18),
            ('alice', 'A', 19)]
```

Documentation for `sorted()` may be found at: https://docs.python.org/3/library/functions.html#sorted.

**Q1:** What is returned when `sorted()` is invoked on the `students` list?  (E.g., when we run `sorted(students)`.)

In this case, we are sorting a list whose elements are tuples.  Describe how the algorithm used by `sorted()` works in the case that two tuples are equal in their first element.  (You can find this information in the Python documentation, or by trying it out yourself.)

<!-- grade_id: cell-2d58e657adf85d40, points: 2, grade: true, solution: true -->
The sorted() function sort the list of tuples starting from elements in index 0. Since our first elements are names in strings, the function compares the first letter (at index 0) of each name. The letter than appear first in the alphabet is considered smaller and therefore placed first. When the first element in two tuples are equal, then the function moves on to the elements at index 1 (grade). 

**Q2:** Look at the documentation for `sorted()`.  If we wanted to sort the `students` list based on the second element of the tuple (the grade), we want a function that is used to extract the "key" from a list element (tuple) to retrieve that second element.  Note that this key is similar, but not quite the same, as a key in a dictionary.

Write a function

```
    getGrade(record)
```

that, given a student record that consists of a 3-tuple, accesses and returns the middle element of the tuple.  You are welcome to define `getGrade` as a lambda function.  With this `getGrade` function object passed as the `key` parameter, call `sorted()` on the list `students` and assign the result to the variable `gradeSorted`.

<!-- grade_id: cell-06893f6918eb7e27, points: 0, grade: false, solution: true -->
```python
def getGrade(record): 
    """
    Returns the grade of input student.
    """
    return record[1]

gradeSorted = sorted(students, key= getGrade)

```

```python
# Printing cell
print(gradeSorted)
```

Output:
```
[('chris', 'A', 18), ('alice', 'A', 19), ('bob', 'B', 19), ('chris', 'B', 18)]
```

<!-- grade_id: cell-0323a68fb56e401b, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert getGrade(('alice', 'A', 18)) == 'A'
assert getGrade(('chris', 'B', 18)) == 'B'

assert isinstance(gradeSorted, list)
assert gradeSorted[0] == ('chris', 'A', 18)
```

**Q3:** Consider your result: when there are ties in the "key", do you get the elements involved in the tie in the same order as in the original list, or in different orders?

<!-- grade_id: cell-d2d1f9d66c965c50, points: 1, grade: true, solution: true -->
The order is same as the original list. So what element comes first in the list is placed first in the sorted list. 

---

## Part B: Returning a function

**Q4:** Write a function

```
    whichIsBigger(f1,f2,n)
```

which computes `f1(n)` and `f2(n)`, and returns whichever _function_ (`f1` or `f2`) results in the larger value when an integer `n` is passed as a parameter. In case of a tie, return `f1`.

Test with `f1(x)` defined as $x^3$ and `f2(x)` defined as $5x^2+4$ and using $n=2$. Test again with $n=100$.

<!-- grade_id: cell-76bddeeaf8ffa2ff, points: 0, grade: false, solution: true -->
```python
def whichIsBigger(f1,f2,n):
    """
    Computes f1(n) and f2(n), and returns whichever of f1 or f2 gives a larger result.
    Returns f1 in the case of a tie.
    """
    num1 = f1(n)
    num2 = f2(n) 
    if num1 >= num2: 
        return f1
    else:
        return f2
```

```python
# Printing cell

def f1(x):
    return x**3

def f2(x):
    return 5*x*x + 4

print(whichIsBigger(f1,f2,2))
print(whichIsBigger(f1,f2,100))
```

Output:
```
<function f2 at 0x10c16a5f0>
<function f1 at 0x10c1df1c0>
```

<!-- grade_id: cell-95a4c27fc69bb970, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert whichIsBigger.__doc__ is not None # don't remove the docstring!

def f1(x):
    return x**3 - 7

def f2(x):
    return 5*x*x + 4*x + 2

assert whichIsBigger(f1,f2,0) == f2
assert whichIsBigger(f1,f2,-10) == f2
assert whichIsBigger(f1,f2,2) == f2
assert whichIsBigger(f1,f2,50) == f1
```

---

## Part C: Lambda expressions

**Q5:** Write a lambda expression an assign it to the variable `greetPerson` such that it returns a greeting when invoked with a single argument.

For example, the expression `greetPerson("Eliza")` should result in the string `Hello Eliza!`, and `greetPerson("Robert B.")` should result in `Hello Robert B.!".

<!-- grade_id: cell-260b5598d856ee88, points: 0, grade: false, solution: true -->
```python
greetPerson = lambda x: "Hello " + str(x) + "!"
```

```python
# Printing cell
print(greetPerson("Eliza"))
print(greetPerson("Robert B."))
```

Output:
```
Hello Eliza!
Hello Robert B.!
```

```python
# Testing cell
assert(greetPerson("Eliza") == "Hello Eliza!")
assert(greetPerson("Robert B.") == "Hello Robert B.!")
assert(greetPerson("everyone") == "Hello everyone!")
assert(greetPerson("") == "Hello !")
```

**Q6:** Write a lambda function that takes three parameters and multiplies them together.  Assign the lambda function expression to `mult3`.

<!-- grade_id: cell-2f45e6d4eb6451d7, points: 0, grade: false, solution: true -->
```python
mult3 = lambda x,y,z: x*y*z
```

```python
# Printing cell
mult3(10, 12, 9)
```

Output:
```
1080
```

<!-- grade_id: cell-2657c0facf794201, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert mult3(10, 12, 9) == 1080
assert mult3(10, -12, 9) == -1080
```

**Q7:** Write a lambda expression and assign it to the variable `evenSquared` such that it squares the value of its single argument if the argument is even, and *cubes* the argument if it is odd.

Afterward, for example, expression `evenSquared(5)` should yield `125`, while `evenSquared(4)` should yield `16`.

Hint: We need an *expression* to handle an `if` type condition.

```python
evenSquared = lambda x: (x**2) if x % 2 == 0 else (x**3)
```

```python
# Testing cell
print(evenSquared(5))
print(evenSquared(4))
```

Output:
```
125
16
```

<!-- grade_id: cell-a34541c503ff5a43, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert(evenSquared(1)==1)
assert(evenSquared(2)==4)
assert(evenSquared(3)==27)
assert(evenSquared(4)==16)
assert(evenSquared(0)==0)
```

```python

```

---

## Part D: `map`

**Mapping Functions**

Python has a built-in function `map()` 

- https://docs.python.org/3/library/functions.html#map

whose parameters are a function object followed by one or more iterables.  If there are only two arguments, the provided first argument function (which should take a single argument) is invoked over the elements of the second argument.  This gives us a built-in way to apply a unary function element-wise over a list or other sequence.

If we pass, as arguments to `map()`, a function followed by *two* iterables, then the function is expected to have *two* parameters, and the function is called for each pair of elements coming from the first and second iterable.  This gives us a built-in way to apply a binary function element-wise over two lists/vectors/sequences.

Note that, for efficiency reasons, `map()` actually returns an *iterator* that, like `range()`, can be used as the sequence in a `for` statement, and waits to compute elements until needed.  If we want a *list* instead of an iterator, we must pass the iterator as the argument to the `list()` function.

**Q8:** Write a function

```
    incrementAge(record)
```

that, given a 3-tuple, `record`, of a student, composed of a name, a grade, and an age, returns a new tuple with the same name and grade, but with the age incremented by one.

Implement this function as a _lambda function_.

<!-- grade_id: cell-9b18fbad2b616883, points: 0, grade: false, solution: true -->
```python
incrementAge = lambda record: (record[0], record[1], record[2]+1)
```

<!-- grade_id: cell-f8e7dee0bb2c7605, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert incrementAge(('chris', 'A', 18)) == ('chris', 'A', 19)
```

**Q9:** Use `map()` to apply the `incrementAge` function to the `students` list, and create a list assigned to variable `students2`.

<!-- grade_id: cell-17560d667aafabb2, points: 0, grade: false, solution: true -->
```python
students = [('bob',   'B', 19),
            ('chris', 'B', 18),
            ('chris', 'A', 18),
            ('alice', 'A', 19)]

new_age =[]
for student in students: 
    new = incrementAge(student)
    new_age.append(new)
students2 = new_age
```

```python
# Printing cell
print(students2) # should print a list, not a "map" object
```

Output:
```
[('bob', 'B', 20), ('chris', 'B', 19), ('chris', 'A', 19), ('alice', 'A', 20)]
```

<!-- grade_id: cell-a1e0506a43e4b068, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert isinstance(students2, list)
assert len(students2) == 4
assert students2[1] == ('chris', 'B', 19)
assert students2[-1] == ('alice', 'A', 20)
```

**Q10:** Write a function

```
    vector_mult(L1, L2)
```

that performs element-wise multiplication between the elements in `L1` and `L2`.  You may assume that `L1` and `L2` are the same length.  You *must* use `map()` in the `vector_mult` function, although you may write a helper function.  Use lambda functions where you can (i.e. where the function can be formed as an expression).

<!-- grade_id: cell-8c8cf63b35739800, points: 0, grade: false, solution: true -->
```python
def vector_mult(L1,L2): 
    """
    Takes in two lists of vectors and returns a list of element-wise multiplication between their elements, assuming L1 and L2 are the same length.
    """
    hi = map(lambda x,y: x*y, L1, L2)
    return list(hi)
```

```python
# Printing cell
A = [1, 2, 3]
B = [5, 10, 1]
print(vector_mult(A, B)) # should be [5, 20, 3]
```

Output:
```
[5, 20, 3]
```

<!-- grade_id: cell-685f2e739829a093, points: 2, grade: true, solution: false -->
```python
# Testing cell
A = [5, 3, 1, 2]
B = [2, 1, 4, 3]
assert vector_mult(A, B) == [10, 3, 4, 6]
```