# Denison DA210/CS181 SW Lab #6 - Step 1

Before you get your checkpoints, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import pandas as pd

datadir = "publicdata"
```

---

## Part A

**Q1:** Consider the data on religions and income, gathered by Pew Research Center and hosted at this link: https://github.com/chendaniely/pandas_for_everyone/blob/master/data/pew.csv.

The data is also available as `pew.csv` in the data folder.  In the code cell that follows, read the data into a DataFrame assigned to `df_rel_untidy`.

<!-- grade_id: cell-cefd70c125c26355, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display part of the DataFrame
df_rel_untidy.head()
```

<!-- grade_id: cell-42c6a8c97256fdaa, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert df_rel_untidy.shape == (18,11)
assert df_rel_untidy.iloc[0,0] == "Agnostic"
assert df_rel_untidy.iloc[0,1] == 27
assert df_rel_untidy.iloc[5,0] == "Evangelical Prot"
assert df_rel_untidy.columns[5] == "$40-50k"
assert df_rel_untidy.iloc[14,8] == 14
```

> You've reached the first checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 1: What are the independent variable(s) and dependent variable(s) for this dataset?  Give one of the three tidy-data requirements that this dataset fails to satisfy and explain why.

**Q2:** Transform as needed to make it tidy.  Store the normalized (tidy) data frame as `df_rel`.

<!-- grade_id: cell-63953f65aa520b07, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# Display part of the tidy DataFrame
df_rel.head()
```

<!-- grade_id: cell-42c6a8c97256fdaa, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert df_rel.shape == (180,3)
assert df_rel.iloc[0,0] == "Agnostic"
assert df_rel.iloc[0,1] == "<$10k"
assert df_rel.iloc[0,2] == 27
assert df_rel.iloc[41,0] == "Evangelical Prot"
assert df_rel.iloc[89,1] == "$40-50k"
assert df_rel.iloc[104,2] == 14
```

---

## Part B

**Q3:** Read the data in `pop_columns.csv` into a `DataFrame` called `pop_untidy`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the untidy DataFrame
pop_untidy
```

> You've reached the second checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 2: What are the independent variable(s) and dependent variable(s) for this dataset?  Would either `melt` or `pivot` be useful to make this data tidy?

**Q4:** This dataset fails to satisfy `TidyData1`: each column represents the `population` for a different value of `year`.  Make this dataset tidy.

Your result should use both the `code` and `year` as a multi-level index and have one data column (`pop`), and twelve rows.  Note that you can use `sort_index` to reorganize your data after setting the index.  Store the resulting `DataFrame` in a variable `pop_tidy`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the tidy DataFrame
pop_tidy
```

```python
# Testing cell
assert pop_tidy.shape == (12, 1)
assert pop_tidy.loc[("USA", 2015), "pop"] == 320.74
```

---

## Part C

**Q5:** Read the data in `popgdp_columns.csv` into a `DataFrame` called `ind_untidy`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the untidy DataFrame
ind_untidy
```

> You've reached the third and final checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 3: What are the independent variable(s) and dependent variable(s) for this dataset?  What steps would you take to make this data tidy?

**Q6:** \[Optional\] This dataset fails to satisfy `TidyData1`: each column represents either the `population` or `GDP` for a different value of `year`.

Make this dataset tidy.  Your result should a multi-level row index of `code` and `year`, two data columns, and twelve rows.  Store the resulting `DataFrame` in a variable `ind_tidy`.

```python
# YOUR CODE HERE
raise NotImplementedError()

# Display the tidy DataFrame
ind_tidy
```

```python
# Testing cell
assert ind_tidy.shape == (12, 2)
assert ind_tidy.loc[("USA", 2015), "pop"] == 320.74
assert ind_tidy.loc[("GBR", 2017), "gdp"] == 2637.87
```

---

---

## Part D

How much time (in minutes/hours) did you spend on this lab outside of class?

YOUR ANSWER HERE