# Denison DA210/CS181 Homework 1.c - Step 2

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os

datadir = "publicdata"
```

---

## Part C

**Q5:** Write a function
```
    lineLengths(filepath)
```
that processes a file, line by line, and accumulates and returns a list of tuples, one per line.  Each tuple consists of the line number and the length of the line, excluding any leading or trailing whitespace (spaces, tabs, or newlines).  Line numbers in a file start at 1.  So, for the `hello.txt` file:

    Hello, world!
    Welcome to 'Introduction to Data Systems'.

the result is the list of tuples: `[(1, 13), (2, 42)]`.

<!-- grade_id: cell-c5e9ddfe460baaf5, points: 0, grade: false, solution: true -->
```python
def lineLengths(filepath): 
    """
    Returning a list of tuples containing the line number and number of words in that line. 
    """
    with open(filepath, "r") as word:
        num_line =1
        lst =[]
        for line in word:
            line = line.strip()
            lst.append((num_line, len(line)))
            num_line= num_line+1
        return lst
```

<!-- grade_id: cell-17559d5c71f8dc5b, points: 3, grade: true, solution: false -->
```python
filepath = os.path.join(datadir, "hello.txt")

assert lineLengths.__doc__ is not None # don't forget the docstring

# Sanity check
print(lineLengths(filepath))

# Run tests
assert lineLengths(filepath) == [(1, 13), (2, 42)]

filepath = os.path.join(datadir, "tennyson.txt")
assert lineLengths(filepath) == [(1, 26), (2, 24), (3, 0), (4, 40), (5, 14)] 
```

Output:
```
[(1, 13), (2, 42)]
```

**Q6:** Consider a variation of the babynames and counts file as depicted below:

    22127,      Jacob
    18002,      Ethan
    17350,    Michael
    17179,     Jayden
    17051,    William
    16756,  Alexander

Each line of the file captures one data case/observation.  The values are separated by commas, with the count occurring first and the name second, and spaces are used to align the columns of data to make it easier for a human reader.

Write a function

```
    readNamesCounts(filepath)
```

that processes the `filepath` file and yields a tuple whose first element is a reference to a list of names, and whose second element is a reference to a list of **integer** counts.

<!-- grade_id: cell-a10993da7de47401, points: 0, grade: false, solution: true -->
```python
def readNamesCounts(filepath): 
    """
    Creating tuples from filepath in which the first element is a list of names and the second elements is a list of counts. 
    """
    with open(filepath, "r") as file: 
        nameList = []
        countList = []
        for line in file: 
            line = line.split(",")
            nameList.append(line[1])
            countList.append(int(line[0])) #fixed int
        return nameList, countList
            
```

<!-- grade_id: cell-dba242eb21fd0162, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert readNamesCounts.__doc__ is not None # don't forget the docstring

path = os.path.join(datadir, "babynames.txt")
assert os.path.isfile(path)

namelist, countlist = readNamesCounts(path)
assert len(namelist) == 6
assert len(countlist) == 6
assert "Jacob" in namelist[0]
```

---

---
## Part D

How much time (in minutes/hours) did you spend on this homework assignment?

20 minutes

Who was your partner for this assignment?  If you worked alone, say so instead.

Alone