# Denison DA210/CS181 Homework 4.b - Step 2

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

#### Import Python modules and load "SQL Magic"

```python
import pandas as pd
import os
import os.path
import json
import sys
import importlib

module_dir = "../../modules"
module_path = os.path.abspath(module_dir)
if not module_path in sys.path:
    sys.path.append(module_path)

%load_ext sql
```

```python
def getsqlite_creds(dirname=".",filename="creds.json",source="sqlite"):
    """ Using directory and filename parameters, open a credentials file
        and obtain the two parts needed for a connection string to
        a local provider using the "sqlite" dictionary within
        an outer dictionary.  
        
        Return a scheme and a dbfile
    """
    assert os.path.isfile(os.path.join(dirname, filename))
    with open(os.path.join(dirname, filename)) as f:
        D = json.load(f)
    sqlite = D[source]
    return sqlite["scheme"], sqlite["dbdir"], sqlite["database"]
```

```python
scheme, dbdir, database = getsqlite_creds(source="sqlite1")
template = '{}:///{}/{}.db'
cstring = template.format(scheme, dbdir, database)
print("Connection string:", cstring)
```

#### Establish Connection from Client to Server

```python
%sql $cstring
```

---

## Part B: SQL Group By Operations

In the following cells, your only action is to remove the `# YOUR CODE HERE` and `raise NotImplementedError()` lines, and then put a valid SQL statement as the **value** of string variable `query`.

In each case, when you execute the cell, the query will be sent to the database management system, a result obtained, and the result converted into a `pandas` data frame, whose prefix is shown.  (This allows for testing the results as well as displaying them in your Jupyter Notebook.)

_Note: you may want to view the `book` database in SQLiteStudio to get an idea of what the field names are as you work through these exercises._

**Q5:** Using the SQL table `countries` in the `book` database, use a select query to answer the question: how many countries are there in each region?  Select the `region` field as well as a new column, `new`, which answers this question.

<!-- grade_id: cell-4c8f2803414fa403, points: 0, grade: false, solution: true -->
```python
query5 = """
"""
# YOUR CODE HERE
raise NotImplementedError()

resultset5 = %sql $query5
resultdf5 = resultset5.DataFrame()
resultdf5
```

<!-- grade_id: cell-ab373181b399dce2, points: 0, grade: true, solution: false -->
```python
# Testing cell
assert len(resultdf5) == 7
assert len(resultdf5.columns) == 2
assert 58 in list(resultdf5['new'])
assert resultdf5.set_index("region").loc["North America", "new"] == 3
```

**Q6:** Use the `indicators` table in the `book` database to find the total world population in each year (as the sum of country populations). Use the alias `total_pop` for your new column, and select the `year` and `total_pop` fields.  Sort the result in descending `year` value.

<!-- grade_id: cell-4e00b362912f0743, points: 0, grade: false, solution: true -->
```python
query6 = """
"""
# YOUR CODE HERE
raise NotImplementedError()

resultset6 = %sql $query6
resultdf6 = resultset6.DataFrame()
resultdf6.head()
```

<!-- grade_id: cell-9e2f0240bddcc016, points: 0, grade: true, solution: false -->
```python
# Testing cell
assert len(resultdf6) == 59
assert resultdf6.loc[0,'total_pop'] > 7500
assert resultdf6.loc[0,'total_pop'] < 7600
assert resultdf6.loc[58,'total_pop'] > 3014
assert resultdf6.loc[58,'total_pop'] < 3015
```

**Q7:** Treating your query above as a subquery (without the ORDER BY), find the minimum for `total_pop` over all years. Use the alias `m` for the new column.

<!-- grade_id: cell-b96de6baf5f97fad, points: 0, grade: false, solution: true -->
```python
query7 = """
"""
# YOUR CODE HERE
raise NotImplementedError()

resultset7 = %sql $query7
resultdf7 = resultset7.DataFrame()
resultdf7.head()
```

<!-- grade_id: cell-90e29bdba973d83c, points: 0, grade: true, solution: false -->
```python
# Testing cell
assert len(resultdf7) == 1
assert resultdf7.loc[0,'m'] > 3014
assert resultdf7.loc[0,'m'] < 3015
```

**Q8:** Not all countries are growing, so the largest population a country ever had might be in a previous year. For each country code in `indicators`, find the max population that country ever had. Alias your new column as `max_pop`. You should have one row per country, and select the fields `code` and `max_pop`. Don't change the order (that is, your records should still be ordered by `code`).

<!-- grade_id: cell-a872397ffe55566d, points: 0, grade: false, solution: true -->
```python
query8 = """
"""
# YOUR CODE HERE
raise NotImplementedError()

resultset8 = %sql $query8
resultdf8 = resultset8.DataFrame()
resultdf8.head()
```

<!-- grade_id: cell-dfa099931224150d, points: 0, grade: true, solution: false -->
```python
# Testing cell
assert len(resultdf8) == 218
assert resultdf8.loc[0,'max_pop'] == 0.11
assert resultdf8.loc[1,'max_pop'] == 37.17
```

**Q9:** Copy your query from the previous question, and then modify it to find all records where the max population is less than 1 (remember, this is measured in millions of people).  Use a `HAVING` clause.  Keep the original ordering of the data (alphabetically, by `code`).

<!-- grade_id: cell-6e74056737f411c0, points: 0, grade: false, solution: true -->
```python
query9 = """
"""
# YOUR CODE HERE
raise NotImplementedError()

resultset9 = %sql $query9
resultdf9 = resultset9.DataFrame()
resultdf9.head()
```

<!-- grade_id: cell-65bcd17ea22c6762, points: 0, grade: true, solution: false -->
```python
# Testing cell
assert len(resultdf9) == 58
assert resultdf9.loc[0,'max_pop'] == 0.11
assert resultdf9.loc[1,'max_pop'] == 0.08
```

**Q10:** Using the `indicators` table in the `book` database, find the total world population in each year (as the sum of country populations), then return the rows where the total population is greater than 6000 (measured in millions of people).  Use the alias `total_pop` for your new column.

<!-- grade_id: cell-bee73edca4961f80, points: 0, grade: false, solution: true -->
```python
query10 = """
"""
# YOUR CODE HERE
raise NotImplementedError()

resultset10 = %sql $query10
resultdf10 = resultset10.DataFrame()
resultdf10.head()
```

<!-- grade_id: cell-5b9767551411a51a, points: 0, grade: true, solution: false -->
```python
# Testing cell
assert len(resultdf10) == 20
assert resultdf10.loc[0,'total_pop'] < 6014
assert resultdf10.loc[1,'total_pop'] > 6090
```

---

---

## Part C

**Q11:** How much time (in minutes/hours) did you spend on this homework assignment?

YOUR ANSWER HERE

**Q12:** Who was your partner for this assignment?  If you worked alone, say so instead.

YOUR ANSWER HERE