# Denison DA210/CS181 Homework 2.e - Step 1

Before you turn this problem in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells** (in the menubar, select Kernel$\rightarrow$Restart And Run All).

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import pandas as pd

datadir = "publicdata"
```

---

## Part A: Concatenating along the row dimension

**Q1:** In the data directory, you will find two csv files, `educationTop.csv` and `educationBottom.csv`, both based on data hosted by `www.census.gov`.  Both have the same columns, and each row is a U.S. metropolitan area, with data on population, education, and unemployment.  The first `csv` file contains metropolitan areas starting with A-K, and the second starting with L-Z.

Read both into `pandas` data frames `education_top` and `education_bottom`, using the column `GEO.id2` as an index.

Be careful, as these files have **two** lines that might be considered header lines.  The first one is more descriptive, and the second one is more concise and gives succint variable names.  So when we read from the csv, we want to **skip** the first line and then use the second line as the header line.  The `pandas` `read_csv` constructor can do this for you, but you have to discover and use the correct arguments.  (Documentation: https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html)

<!-- grade_id: cell-b25fab8c12339430, points: 0, grade: false, solution: true -->
```python
filepath1 = os.path.join(datadir, "educationTop.csv")
file1 = pd.read_csv(filepath1, skiprows=[0], index_col = ["GEO.id2"]) 
education_top = pd.DataFrame(file1)

filepath2 = os.path.join(datadir, "educationBottom.csv")
file2 = pd.read_csv(filepath2, skiprows=[0], index_col = ["GEO.id2"]) 
education_bottom = pd.DataFrame(file2)
```

```python
# Display a sample of the "top" DataFrame
education_top.head(2)
```

Output:
```
GEO.id       GEO.display-label  HD01_VD01  HD02_VD01  \
GEO.id2                                                                 
10180    310M100US10180  Abilene, TX Metro Area      82270       1582   
10420    310M100US10420    Akron, OH Metro Area     369311       1080   

         HD01_VD02  HD02_VD02  HD01_VD03  HD02_VD03  HD01_VD04  HD02_VD04  \
GEO.id2                                                                     
10180        13820       1873       6484       1272          0        206   
10420        26613       2972      14519       2048          0        182   

         ...  HD01_VD25  HD02_VD25  HD01_VD26  HD02_VD26  HD01_VD27  \
GEO.id2  ...                                                          
10180    ...        579        347      12909       1564      12770   
10420    ...         43         70      97693       4241      94492   

         HD02_VD27  HD01_VD28  HD02_VD28  HD01_VD29  HD02_VD29  
GEO.id2                                                         

```

```python
# Display a sample of the "bottom" DataFrame
education_bottom.head(2)
```

Output:
```
GEO.id            GEO.display-label  HD01_VD01  HD02_VD01  \
GEO.id2                                                                      
29100    310M100US29100  La Crosse, WI-MN Metro Area      67151        664   
29140    310M100US29140     Lafayette, IN Metro Area      94481       1352   

         HD01_VD02  HD02_VD02  HD01_VD03  HD02_VD03  HD01_VD04  HD02_VD04  \
GEO.id2                                                                     
29100         2114        654       1202        471          0        184   
29140         8844       1537       5463       1334          0        184   

         ...  HD01_VD25  HD02_VD25  HD01_VD26  HD02_VD26  HD01_VD27  \
GEO.id2  ...                                                          
29100    ...          0        184      17955       1585      17592   
29140    ...         19         31      27825       2372      26749   

         HD02_VD27  HD01_VD28  HD02_VD28  HD01_VD29  HD02_VD29  
GEO.id2                                      
```

<!-- grade_id: cell-8f07ee547e7c1ff8, points: 3, grade: true, solution: false -->
```python
# Testing Cell
assert education_top.shape == (176, 60)
assert education_bottom.shape == (190, 60)
assert education_top.index.name == "GEO.id2"
assert education_bottom.index.name == "GEO.id2"
assert "HD01_VD01" in education_top
assert "HD01_VD01" in education_bottom
assert education_top.iloc[8,2] == 129504
assert education_bottom.iloc[8,4] == 39259
```

**Q2:** Concatenate these two data frames along the row dimension (with the top one on top), and call the result `educationDF`.

```python
educationDF = pd.concat([education_top, education_bottom], axis = 0)

# Display a sample of the DataFrame
educationDF.head(4)
```

Output:
```
GEO.id                       GEO.display-label  HD01_VD01  \
GEO.id2                                                                      
10180    310M100US10180                  Abilene, TX Metro Area      82270   
10420    310M100US10420                    Akron, OH Metro Area     369311   
10500    310M100US10500                   Albany, GA Metro Area      78159   
10580    310M100US10580  Albany-Schenectady-Troy, NY Metro Area     466236   

         HD02_VD01  HD01_VD02  HD02_VD02  HD01_VD03  HD02_VD03  HD01_VD04  \
GEO.id2                                                                     
10180         1582      13820       1873       6484       1272          0   
10420         1080      26613       2972      14519       2048          0   
10500         1072      13689       2238       7681       1645          0   
10580         1818      29031       3211      16906       2415          0   

         HD02_VD04  ...  HD01_VD25  HD02_VD25  HD01_VD26  HD02_VD26  \
GEO.id2       
```

```python
# Testing Cell
assert educationDF.shape == (366, 60)
assert educationDF.index.name == "GEO.id2"
assert "HD01_VD01" in educationDF
assert educationDF.iloc[8,2] == 129504
assert educationDF.iloc[184,4] == 39259
```

---

## Part B: Concatenating along the column dimension

**Q3:** In the data directory, you will find two csv files, `educationLeft.csv` and `educationRight.csv`, both based on data hosted by `www.census.gov`.  Both have the same rows, and each row is a U.S. metropolitan area, with data on population, education, and unemployment.  The first has information on individuals without a college degree, and the second has information on individuals with a college degree.

Read both into `pandas` data frames `education_left` and `education_right`.  The `GEO.display-label` field should serve as the row index.

The same cautions on the CSV having an informational row of headers followed by the header variable names applies to these two CSV files as well.

```python
filepath3 = os.path.join(datadir, "educationLeft.csv") 
file3 = pd.read_csv(filepath3, skiprows=[0], index_col = ["GEO.display-label"]) 
education_left = pd.DataFrame(file3) 

filepath4 = os.path.join(datadir, "educationRight.csv") 
file4 = pd.read_csv(filepath4, skiprows=[0], index_col = ["GEO.display-label"]) 
education_right = pd.DataFrame(file4) 


```

```python
# Display a sample of the left DataFrame
education_left.head(2)
```

Output:
```
GEO.display-label  HD01_VD01  HD02_VD01  HD01_VD02  HD02_VD02  \
0  Abilene, TX Metro Area      82270       1582      13820       1873   
1    Akron, OH Metro Area     369311       1080      26613       2972   

   HD01_VD03  HD02_VD03  HD01_VD04  HD02_VD04  HD01_VD05  ...  HD01_VD11  \
0       6484       1272          0        206       6484  ...        242   
1      14519       2048          0        182      14519  ...          0   

   HD02_VD11  HD01_VD12  HD02_VD12  HD01_VD13  HD02_VD13  HD01_VD14  \
0        230      16089       2007      15268       2047        821   
1        182      91228       4194      82462       4125       8766   

   HD02_VD14  HD01_VD15  HD02_VD15  
0        437       8729       1335  
1       1850      28828       3019  

[2 rows x 31 columns]
```

```python
# Display a sample of the right DataFrame
education_right.head(2)
```

Output:
```
HD01_VD16  HD02_VD16  HD01_VD17  HD02_VD17  HD01_VD18  \
GEO.display-label                                                               
Abilene, TX Metro Area      26269       2060      20430       2075        705   
Akron, OH Metro Area       111011       4257      86994       3887        123   

                        HD02_VD18  HD01_VD19  HD02_VD19  HD01_VD20  HD02_VD20  \
GEO.display-label                                                               
Abilene, TX Metro Area        307      19725       1978      18421       1902   
Akron, OH Metro Area          119      86871       3903      79172       3642   

                        ...  HD01_VD25  HD02_VD25  HD01_VD26  HD02_VD26  \
GEO.display-label       ...                                               
Abilene, TX Metro Area  ...        579        347      12909       1564   
Akron, OH Metro Area    ...         43         70      97693       4241   

                        HD01_VD27  HD02_VD27  HD01_VD28  HD02_VD28  HD01_
```

<!-- grade_id: cell-f413a9fabbcd0427, points: 3, grade: true, solution: false -->
```python
# Testing Cell
assert education_left.shape == (366, 30)
assert education_right.shape == (366, 28)
assert education_left.index.name == "GEO.display-label"
assert education_right.index.name == "GEO.display-label"
assert "HD01_VD01" in education_left
assert "HD01_VD01" not in education_right
assert "Abilene, TX Metro Area" in education_left.index
assert "Abilene, TX Metro Area" in education_right.index
assert education_left.iloc[10,3] == 2265
assert education_right.iloc[10,2] == 63125
```

**Q4:** Concatenate these two data frames along the column dimension, and call the result `educationDF2`.

```python
educationDF2 = pd.concat([education_left, education_right], axis =1)

# Display a sample of the DataFrame
educationDF2.head(4)
```

Output:
```
HD01_VD01  HD02_VD01  HD01_VD02  \
GEO.display-label                                                         
Abilene, TX Metro Area                      82270       1582      13820   
Akron, OH Metro Area                       369311       1080      26613   
Albany, GA Metro Area                       78159       1072      13689   
Albany-Schenectady-Troy, NY Metro Area     466236       1818      29031   

                                        HD02_VD02  HD01_VD03  HD02_VD03  \
GEO.display-label                                                         
Abilene, TX Metro Area                       1873       6484       1272   
Akron, OH Metro Area                         2972      14519       2048   
Albany, GA Metro Area                        2238       7681       1645   
Albany-Schenectady-Troy, NY Metro Area       3211      16906       2415   

                                        HD01_VD04  HD02_VD04  HD01_VD05  \
GEO.display-label                                              
```

```python
# Testing Cell
assert educationDF2.shape == (366, 58)
assert educationDF2.index.name == "GEO.display-label"
assert "HD01_VD01" in educationDF2
assert educationDF2.iloc[10,3] == 2265
assert educationDF2.iloc[10,32] == 63125
```

---

## Part C: Joining/merging data frames

**Q5:** In the data directory, you will find two csv files, `educationLeftJ.csv` and `educationRightJ.csv`, both based on data hosted by `www.census.gov`.  In both, rows represent U.S. metropolitan area, with data on population, education, and unemployment.  However, they do not have exactly the same set of rows, and the columns are totally different except for the index column `Geography`.

In these CSV files, there is not an "extra" header line.  The row label index should come from `Geography` unless you are instructed otherwise.

Read both into `pandas` data frames, with names `education_left_j` and `education_right_j`.

<!-- grade_id: cell-f3cb3f384a53b7b2, points: 0, grade: false, solution: true -->
```python
left = os.path.join(datadir, "educationLeftJ.csv")
left = pd.read_csv(left, index_col = ["Geography"]) 
education_left_j = pd.DataFrame(left) 


right = os.path.join(datadir, "educationRightJ.csv")
right = pd.read_csv(right, index_col = ["Geography"]) 
education_right_j = pd.DataFrame(right) 


```

```python
# Display a sample of the left_j DataFrame
education_left_j.head(2)
```

Output:
```
Estimate; Less than high school graduate: - In labor force:  \
Geography                                                                             
Abilene, TX Metro Area                                               6484             
Akron, OH Metro Area                                                14519             

                        Estimate; Less than high school graduate: - In labor force: - In Armed Forces  \
Geography                                                                                               
Abilene, TX Metro Area                                                  0                               
Akron, OH Metro Area                                                    0                               

                        Estimate; Less than high school graduate: - In labor force: - Civilian:  \
Geography                                                                                         
Abilene, TX Metro Area                                  
```

```python
# Display a sample of the right_j DataFrame
education_right_j.head(2)
```

Output:
```
Estimate; Some college or associate's degree: - In labor force:  \
Geography                                                                                 
Abilene, TX Metro Area                                              20430                 
Akron, OH Metro Area                                                86994                 

                        Estimate; Some college or associate's degree: - In labor force: - In Armed Forces  \
Geography                                                                                                   
Abilene, TX Metro Area                                                705                                   
Akron, OH Metro Area                                                  123                                   

                        Estimate; Some college or associate's degree: - In labor force: - Civilian:  \
Geography                                                                                             
Abilene, TX Metr
```

<!-- grade_id: cell-a1e03d05287e5c02, points: 4, grade: true, solution: false -->
```python
# Testing Cell
assert education_left_j.shape == (249, 12)
assert education_right_j.shape == (215, 12)
assert education_left_j.index.name == "Geography"
assert education_right_j.index.name == "Geography"
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" in education_left_j
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" not in education_right_j
assert "Estimate; Bachelor's degree or higher: - In labor force:" not in education_left_j
assert "Estimate; Bachelor's degree or higher: - In labor force:" in education_right_j
assert "Battle Creek, MI Metro Area" in education_left_j.index
assert "Battle Creek, MI Metro Area" not in education_right_j.index
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" not in education_left_j.index
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" in education_right_j.index
assert education_left_j.iloc[5,6] == 38994
assert education_right_j.iloc[9,2] == 59784
```

**Q6:** Make a copy of `education_left_j` called `education_left_orig`.  Starting with `education_left_j`, do a left join to bring in the data from `education_right_j`, storing your answer as `df_j1`.

```python
education_left_orig = education_left_j.copy()

df_j1 = education_left_orig.join(education_right_j, how ="left")

# Display a sample of the joined DataFrame
df_j1.head(2)
```

Output:
```
Estimate; Less than high school graduate: - In labor force:  \
Geography                                                                                                   
Abilene, TX Metro Area                                                                     6484             
Akron, OH Metro Area                                                                      14519             
Albany, GA Metro Area                                                                      7681             
Albany-Schenectady-Troy, NY Metro Area                                                    16906             
Albuquerque, NM Metro Area                                                                30718             
...                                                                                         ...             
Worcester, MA Metro Area                                                                  17925             
Yakima, WA Metro Area                                            
```

```python
# Testing Cell
assert education_left_orig.shape == (249, 12)

assert df_j1.shape == (249, 24)
assert df_j1.index.name == "Geography"
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" in df_j1
assert "Estimate; Bachelor's degree or higher: - In labor force:" in df_j1
assert "Battle Creek, MI Metro Area" in df_j1.index
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" not in df_j1.index
assert df_j1.iloc[5,6] == 38994
assert df_j1.iloc[5,14] == 59784
```

**Q7:** Starting with `education_left_orig`, do an inner join to bring in the data from `education_right_j`, storing your answer as `df_j2`.

```python
df_j2 = education_left_orig.join( education_right_j, how = "inner")
# Display a sample of the joined DataFrame
df_j2.head(2)
```

Output:
```
Estimate; Less than high school graduate: - In labor force:  \
Geography                                                                                              
Abilene, TX Metro Area                                                                6484             
Akron, OH Metro Area                                                                 14519             
Albany, GA Metro Area                                                                 7681             
Albany-Schenectady-Troy, NY Metro Area                                               16906             
Albuquerque, NM Metro Area                                                           30718             
...                                                                                    ...             
Wenatchee-East Wenatchee, WA Metro Area                                               7258             
Wheeling, WV-OH Metro Area                                                            3902             
W
```

```python
# Testing Cell
assert df_j2.shape == (152, 24)
assert df_j2.index.name == "Geography"
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" in df_j2
assert "Estimate; Bachelor's degree or higher: - In labor force:" in df_j2
assert "Battle Creek, MI Metro Area" not in df_j2.index
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" not in df_j2.index
assert "Asheville, NC Metro Area" in df_j2.index
assert df_j2.iloc[5,6] == 38994
assert df_j2.iloc[5,14] == 59784
```

**Q8:** We can "merge" two data frames if they do not have a meaningful index, but do share a column name.  You can use this syntax:

```
    merged_df = pd.merge(df1, df2, on=<colname>, how=<merge_type>)
```

where `<colname>` is the name of the column they share (typically a string), and `<merge_type>` is a string like `"left"`, `"inner"`, etc.

Now read the original csv files in as `edu_left` and `edu_right` with no meaningful index. Then, starting from `edu_left`, do an inner merge along the column `Geography`, storing your answer as `df_j3`.

```python
left = os.path.join(datadir, "educationLeftJ.csv")
edu_left = pd.DataFrame(pd.read_csv(left))
right = os.path.join(datadir, "educationRightJ.csv")
edu_right = pd.DataFrame(pd.read_csv(right))

df_j3 = pd.merge(edu_left, edu_right, on = "Geography", how = "inner")

# Display a sample of the merged DataFrame
df_j3.head(2)
```

Output:
```
Geography  \
0  Abilene, TX Metro Area   
1    Akron, OH Metro Area   

   Estimate; Less than high school graduate: - In labor force:  \
0                                               6484             
1                                              14519             

   Estimate; Less than high school graduate: - In labor force: - In Armed Forces  \
0                                                  0                               
1                                                  0                               

   Estimate; Less than high school graduate: - In labor force: - Civilian:  \
0                                               6484                         
1                                              14519                         

   Estimate; Less than high school graduate: - In labor force: - Civilian: - Employed  \
0                                               5347                                    
1                                              11338           
```

```python
# Testing Cell
assert df_j3.shape == (152, 25)
assert df_j3.index.name == None
assert "Estimate; Less than high school graduate: - In labor force: - Civilian: - Unemployed" in df_j3
assert "Estimate; Bachelor's degree or higher: - In labor force:" in df_j3
assert "Battle Creek, MI Metro Area" not in list(df_j3.Geography)
assert "Bridgeport-Stamford-Norwalk, CT Metro Area" not in list(df_j3.Geography)
assert "Asheville, NC Metro Area" in list(df_j3.Geography)
assert df_j3.iloc[5,7] == 38994
assert df_j3.iloc[5,15] == 59784
```

**Q9:** Explain the difference in the number of rows and columns in `df_j1`, `df_j2`, and `df_j3`.  Be specific.

<!-- grade_id: cell-60bd9e8dd53dd5f0, points: 2, grade: true, solution: true -->
df_j1 is a left join data frame with 249 rows and 24 columns. The left data drame was from education_left_orig, so the 249 rows match with the 249 observations from education_left_orig, plus the original 12 columns from education_left_orig and new 12 columns from education_right_j. The 12 new columns will have values (or NaNs) for each row in the education_left_orig index variable "Geography".   

df_j2 is an innter join data frame with 152 rows and 24 columns. This 152 rows are the common values in the indexed "Geography" that appear in both education_left_orig and education_right_j. Then, it pulls all the according values from all columns in both of the data frames to df_j2. df_j2 should not have any NaN values. 

df_j3 is a merge on "Geography" data frame with 152 rows and 25 columns. This data frame is almost identical as df_j2, but it has "Geography" as a data frame column, not index. 

---

---

## Part D

**Q10:** How much time (in minutes/hours) did you spend on this homework assignment?

40 minutes

**Q11:** Who was your partner for this assignment?  If you worked alone, say so instead.

Alone