# Denison DA210/CS181 Homework 3.b - Step 1

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import io
from lxml import etree
import json
import sys
import os.path
import pandas as pd

datadir = "publicdata"
```

---

## Part A

**Q1:** Write a function:
```
    getLocalXML(filename, datadir=".", parser=None)
``` 
that performs the common steps of creating a path from the given `filename` and `datadir` and parses the XML file, using the passed `parser`, if any, and returns the Element at the **root** of the tree.  If `parser` is not passed, the standard `XMLParser` should be used.

If the file is not found, or if the parse is unsuccessful (due to XML not being "well formed"), the function should return `None`. Remember that if a parse is unsuccessful, the `etree` module raises an exception.  That means that you should have a `try` block, and indented within that block, the `parse()` invocation should occur.  The `try` block is followed by an `except Exception as e:` line, and within that, your return `None`.  If no exception is raised, code execution will proceed beyond the `try`/`except` block, and that is where you would return the root of the parsed tree.

<!-- grade_id: cell-81428647baf27b6c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6f7ca02f25f931f2, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert getLocalXML.__doc__ is not None # don't forget the docstring!

myparser = etree.XMLParser(remove_blank_text=True)

wroot = getLocalXML("widombooks.xml", datadir, myparser)
assert len(wroot) == 8
bad = getLocalXML("foo.xml", datadir, myparser)
assert bad == None
bad2 = getLocalXML("bad.xml", datadir)
assert bad2==None
```

---

---

## Part B

**Q2:** How much time (in minutes/hours) did you spend on this homework assignment?

YOUR ANSWER HERE

**Q3:** Who was your partner for this assignment?  If you worked alone, say so instead.

YOUR ANSWER HERE