# Denison DA210/CS181 Homework 1.b - Step 2

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import io
import sys
from contextlib import redirect_stdout

datadir = "publicdata"
```

---

## Part C

**Q10:** Write a function

```
    myFileList(directoryPath)
```
    
that queries the specified directory, `directoryPath`, for the files and folders contained therein and generates and returns a list with an element for each of the non-hidden **files** in the directory.  This means that your list should **not** include any entries for hidden files (that start with a `.` character), nor should it include entries for directories (i.e., folders).

<!-- grade_id: cell-5a2690ad91df7a9f, points: 0, grade: false, solution: true -->
```python
def myFileList(directoryPath): 
    """
    Generates a list of all non-hidden files in a specified directory path.
    """
    if os.path.exists(directoryPath) == False: 
        return "Directory not found"
    else: 
        file_lst = [] 
        all_lst = os.listdir(directoryPath)
        for i in all_lst:
            if os.path.isfile(directoryPath+"/"+i) == True and i[0] != ".": 
                file_lst.append(i)
        return file_lst
```

<!-- grade_id: cell-07f7d7e0c1519cdf, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert myFileList.__doc__ is not None # don't forget the docstring!

hw_1b_filelist = myFileList(".")
assert len(hw_1b_filelist) == 3
assert "filebasic.ipynb" in hw_1b_filelist
assert "README.md" in hw_1b_filelist
```

Output:
```
['filebasic.ipynb', 'README.md', 'publicdata', 'pythonbasic1.ipynb']
```

<!-- grade_id: cell-4379c9eebdfdaf7e, points: 2, grade: true, solution: false -->
```python
# Testing cell with tougher test cases

publicdata_filelist = myFileList(datadir)
print(publicdata_filelist)
assert "tennyson.txt" in publicdata_filelist
assert "file1" in publicdata_filelist
```

Output:
```
['tennyson.txt', 'moredata', 'twolines.txt', 'names.json', 'file1', '.file2']
['tennyson.txt', 'twolines.txt', 'names.json', 'file1']
```

---

---
## Part D

How much time (in minutes/hours) did you spend on this homework assignment?

more than 1 hour

Who was your partner for this assignment?  If you worked alone, say so instead.

Alone