# Denison DA210/CS181 Homework 1.b - Step 1

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import io
import sys
from contextlib import redirect_stdout
```

## Instructions

- Complete all the problems below by replacing the two-line sequence of `# YOUR CODE HERE` and `raise NotImplementedError()` with your solution code.
- All functions that you define for homework **must have docstrings**.

---
## Part A

**Q1:** Write a function `sum_of_evens` that only adds up the even numbers between 0 and the number $n$, inclusive.

<!-- grade_id: cell-9416b535d19fc949, points: 0, grade: false, solution: true -->
```python
def sum_of_evens(n):
    """
    Compute the sum of the even numbers from 0 to n.
    """
    s = 0
    for i in range(0, n+1):
        if i % 2 == 0:
            s = s + i
    return s
```

<!-- grade_id: cell-d9c5a7f8882a8e91, points: 1, grade: true, solution: false -->
```python
assert sum_of_evens.__doc__ is not None # make sure to have a docstring

# Check that sum_of_evens returns the correct answer for various inputs.
assert sum_of_evens(1) == 0
assert sum_of_evens(5) == 6
assert sum_of_evens(10) == 30
assert sum_of_evens(11) == 30
```

**Q2:** Describe a **different** way to implement `sum_of_evens()`.

<!-- grade_id: cell-0083b8d23a14fc37, points: 1, grade: true, solution: true -->
We can make a list of all numbers from 0 to n, and then add all numbers with odd indicies. 

**Q3:** Use three *different* ways to create three lists with the integer numbers 1 through 10, inclusive.  Use the variable names `myList1`, `myList2`, and `myList3` to hold the results.

<!-- grade_id: cell-38d713dddb2b5cfe, points: 0, grade: false, solution: true -->
```python
def myList1():
    """ 
    Creating a list with all integers from 1 to 10, using for loop.
    """
    l = []
    for i in range(1, 11):
        l.append(i)
    return l

def myList2():
    """
    Creating a list with all integers from 1 to 10, using while loop. 
    """
    l = []
    i = 1
    while i <=10:
        l.append(i)
        i = i +1
    return l

def myList3():
    """
    Creating a list with all integers from 1 to 10, using while loop with the len function. 
    """
    l = []
    i = 1
    while len(l) <10:
        l.append(i)
        i = i + 1
    return l 

myList1 = myList1() 
myList2 = myList2() 
myList3 = myList3()
```

```python
# Testing cell
assert myList1 == [1,2,3,4,5,6,7,8,9,10]
assert myList2 == [1,2,3,4,5,6,7,8,9,10]
assert myList3 == [1,2,3,4,5,6,7,8,9,10]
```

**Q4:** Write a function:

```
    mean(data)
```
 
that takes a list of numbers given by `data` and computes and returns the arithmetic mean of the values in the list.  Your function should work properly even if `data` is an empty list, and should return 0 in that case.

<!-- grade_id: cell-a32002531e280bd3, points: 0, grade: false, solution: true -->
```python
def mean(data):
    """
    Computes the mean of all values in a numerical list. 
    """
    s = 0 
    if len(data) == 0:
        return s
    for i in range(len(data)):
        s = s + data[i]
    return s/len(data)
```

<!-- grade_id: cell-ef96cd921bb7b67b, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert mean.__doc__ is not None # make sure to have a docstring

list1 = [3, 4, 7, 2]
assert mean(list1) == 4
list2 = []
assert mean(list2) == 0
```

---
## Part B

**Q5:** For the following questions, each answer should consist of a single
`print` with a single argument--an expression involving the list `data`, which, for these questions, references the list `[32, 42, 11, 15, 58, 44, 16]`.

Leave the code we give you "as is", and your print should be indented so that your solution becomes the body of the `with`.

Print a list that includes all elements *except* the first element.

<!-- grade_id: cell-8f7be805bddd4762, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    result.write(str(data[1:])+"\n")
print(result.read())

```

<!-- grade_id: cell-69e005cd46307fdb, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "[42, 11, 15, 58, 44, 16]\n"
```

**Q6:** Print the length of the list `data`.

<!-- grade_id: cell-20ff0e0c86a745c9, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    result.write(str(len(data))+"\n")
print(result.read())
```

<!-- grade_id: cell-a1a0a273991fb0a4, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "7\n"
```

**Q7:** Print the third element of the list `data`.

<!-- grade_id: cell-2769c568b2695d86, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    result.write(str(data[2])+"\n")
print(result.read())
```

<!-- grade_id: cell-0f56edb5b71e8344, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "11\n"
```

**Q8:** Print the last element of the list `data`.

<!-- grade_id: cell-30bbcffba0b8be8d, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    result.write(str(data[-1])+"\n")
print(result.read())
```

<!-- grade_id: cell-35bde65789e0106e, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "16\n"
```

**Q9:** Print a list containing the last four elements of `data`.

<!-- grade_id: cell-ff997c6b4a2fc5eb, points: 0, grade: false, solution: true -->
```python
data = [32, 42, 11, 15, 58, 44, 16]

result = io.StringIO()
with redirect_stdout(result):
    result.write(str(data[-4:])+"\n")
print(result.read())

```

<!-- grade_id: cell-b59a34979e5d93d7, points: 1, grade: true, solution: false -->
```python
# Testing cell
assert result.getvalue() == "[15, 58, 44, 16]\n"
```