# Denison DA210/CS181 SW Lab #7 - Step 1

Before you get your checkpoints, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
import json
import pandas as pd

datadir = "publicdata"
```

---

## `json` module operations

Recall the following operations provided by the `json` module in Python:

Function                |  Description
------------------------|----------------------------------------------------
`json.load(file)`       | Read and return the JSON-formatted data structure from the file object `file`.
`json.dump(data, file)` | Write the data structure `data` to the file object `file` in JSON text format (returns `None`).
`json.loads(s)`         | Using the JSON-formatted string given by `s`, interpret and construct and return the corresponding data structure.
`json.dumps(data)`      | Translate the data structure `data` into a JSON-formatted string and return the string.

Note that, when we `dump()` or `dumps()`, we are **converting** from an in-memory data structure that consists of dictionaries, lists, strings, and numbers (that one can "do math" on), into what is fundamentally a **text string**, either referenced by a string variable, or that is now the text contents of a text file in the file system, which could be opened with an editor independently of anything Python'esque.

When we `load()` or `loads()`, we are going the other direction, and are converting from a string or from the contents of a text file in the file system, and are building an in-memory data structure that we can then traverse and compute with.

---

## Part A: Writing to JSON

Typical steps when we want to **create** a file with the JSON text representation of a data structure from our Python program (to be able to send to another scientist, much like we might build a 2D structure and then want to send a CSV file).

1. Create the data structure, making sure it is limited to dictionaries, lists, strings, integers, and floating point numbers in a single data structure.  Call this `data`.
2. Using the built-in Python function `open`, create and open a file for **writing**, specifying the path and name for the desired file.  We can call the file object `file`.
3. From the `json` module, invoke the `dump()` function, passing it `data` and `file`.
4. Close `file` (this is done for you if you used `with` to open the file).

---

**Q1:** Use the four step process given above to write to a JSON-formatted file.  

Create a Python data structure consisting of a list of dictionaries representation of the following table, where student names are strings, GPAs are floating point numbers, and (class) years are integers.

student | gpa | year
--------|-----|-------
Jane    | 3.75 | 3
Bill    | 2.85 | 2
Fred    | 3.5  | 3
Mary    | 3.25 | 1

Call the in-memory data structure `LoD`.  For this exercise, the desired destination file should be in the current directory and will be named `students.json`.

<!-- grade_id: cell-86c98fb8736ed7b3, points: 0, grade: false, solution: true -->
```python
# Make sure you don't already have a file named students.json
path = os.path.join(".", "students.json")
if os.path.isfile(path):
    os.remove(path)

LoD = [{"student" :"Jane","gpa": 3.75, "year":3}, 
        {"student" :"Bill","gpa": 2.85, "year":2}, 
        {"student" :"Fred","gpa": 3.5, "year":3}, 
        {"student" :"Mary","gpa": 3.25, "year":1}]

file = open("students.json", "w")
json.dump(LoD, file)
file.close()
```

<!-- grade_id: cell-b3992585ad29a7cd, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert os.path.isfile(path)
with open(path, "r") as F2:
    LoD2 = json.load(F2)
assert isinstance(LoD2, list)
assert len(LoD2) == 4
assert isinstance(LoD2[0], dict)
assert len(LoD2[0]) == 3
assert 'student' in LoD2[0]
assert 'gpa' in LoD2[0]
assert 'year' in LoD2[0]
assert isinstance(LoD[0]['gpa'], float)
assert isinstance(LoD[0]['year'], int)
assert LoD[0]['gpa'] == 3.75
assert LoD[0]['year'] == 3
```

---

## Part B: Reading from JSON

In the data directory is a text file named `eu_covid.json`.  Right click and select `Editor` to open with a simple text editor.  Be patient, as it is a **large** file.  Then use the following cell to print the first 30 lines of the file.

```python
path = os.path.join(datadir, "eu_covid.json")
with open(path, 'r') as covid_file:
    for i in range(30):
        line = covid_file.readline()
        print(line, end='')
        
covid_file.close()
```

Output:
```
{
   "records" : [
      {
         "dateRep" : "21/09/2020",
         "day" : "21",
         "month" : "09",
         "year" : "2020",
         "cases" : 0,
         "deaths" : 0,
         "countriesAndTerritories" : "Afghanistan",
         "geoId" : "AF",
         "countryterritoryCode" : "AFG",
         "popData2019" : 38041757,
         "continentExp" : "Asia",
         "Cumulative_number_for_14_days_of_COVID-19_cases_per_100000" : "1.69813397"
      },
      {
         "dateRep" : "20/09/2020",
         "day" : "20",
         "month" : "09",
         "year" : "2020",
         "cases" : 125,
         "deaths" : 4,
         "countriesAndTerritories" : "Afghanistan",
         "geoId" : "AF",
         "countryterritoryCode" : "AFG",
         "popData2019" : 38041757,
         "continentExp" : "Asia",
         "Cumulative_number_for_14_days_of_COVID-19_cases_per_100000" : "1.89265706"
      },
```

**Q2:** To be able to process tree-structured data, you first must **understand** the structure before you attempt to process it.  The next couple of sub-questions begin that process based on **inspection** of the text file, and before we convert it into an in-memory data structure.

**A:** What is the JSON data type for the structure at the root of the tree?

**B:** What is the JSON data type for the **value** that the top-level child **maps to**.

**C.** Within the data type you answered for **B**, what is the JSON data type for the elements at this next level of the tree?

<!-- grade_id: cell-b12276de6bbde4dc, points: 3, grade: true, solution: true -->
A. The root of the tree is the dictionary.

B. lists. 

C. dictionaries



> You've reached the first checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 1: What are your answers to sub-questions A-C above?

**Q3:** Create a variable for the path to the `eu_covid.json` JSON file and then `open()` and `load()` from the text file into an in-memory data structure in Python referred to by variable `covid_data`.

<!-- grade_id: cell-f267174d30854f0e, points: 2, grade: true, solution: true -->
```python
path = os.path.join(datadir, "eu_covid.json")
file = open(path, "r")
covid_data = json.load(file)
file.close()
```

**Q4:** If **you** were writing the testing `assert` statements that made sure that `covid_data` contained the in-memory data structure and that your answers to **A**, **B**, and **C** of **Q3** were correct, you would have three lines that checked 

1. the Python data type of the top level, 
2. the Python data type of the child mapping from the root top level to its value, and
3. the Python data type of an element (probably the first) in the elements within the structure asserted by the previous step.

All three of these would take the form of:
```
    assert isinstance(<expression>, <expected data type>)
```
Note that your answers in **Q2** were in terms of the JSON data types, and these asserts are in terms of Python data types.

Make these three assertions in the cell that follows.

<!-- grade_id: cell-b9b0a8fc3f7b8ff4, points: 3, grade: true, solution: true -->
```python
path = os.path.join(".", "students.json")
with open(path, "r") as file:
    test = json.load(file)
    
file.close()
assert isinstance(test, object)
assert isinstance(test[0], dict)
assert isinstance(test[0]["year"], int)

```

**Q5:** Hopefully, you have determined that the top level of the in-memory data structure is a dictionary, that there is only one child mapping, with  maps the string `"records"` to a list, and that the elements of this list are each a dictionary.

One last sanity check you should perform before writing code to traverse the structure:

- because the file is large, it is difficult to be sure that there is **only** one child mapping in the top level dictionary.  So in the cell that follows, just print the set of dictionary **keys** found in the top level dictionary (converted to a list).   If there is only one element in the list, then we are ready to start building a tabular structure from this tree.

<!-- grade_id: cell-5c13676d1f927f2e, points: 1, grade: true, solution: true -->
```python
path = os.path.join(datadir, "eu_covid.json")
file = open(path, "r")
covid_data = json.load(file)
for i in covid_data:
    print(i)
file.close()


for record, value in covid_data.items():
    print(record)
```

Output:
```
records
records
```

**Q6:** If our goal is one or more `pandas` DataFrame tables of **tidy** data, there are a number of possible solution paths we could pursue.  For now, we are going to assume we want **all** of the innermost dictionary fields as the columns in a `pandas` DataFrame result and will perform normalization _after_ we get this single table.

Given the innermost structure, representing the **rows** of our desired DataFrame is a **dictionary**, one path would be for us to build a **List of Dictionaries** data structure by starting with an empty list and then appending a **copy** of each of the dictionaries within the `covid_data['records']` list.

In the cell that follows, write this loop and build a structure named `LoD`.

<!-- grade_id: cell-76959c470d139c00, points: 0, grade: false, solution: true -->
```python
for record, value in covid_data.items():
    LoD = value

# Display a subset of the LoD
for key,value in LoD[0].items():
    print("{0:>60} {1:>15}".format(key, value))
```

Output:
```
dateRep      21/09/2020
                                                         day              21
                                                       month              09
                                                        year            2020
                                                       cases               0
                                                      deaths               0
                                     countriesAndTerritories     Afghanistan
                                                       geoId              AF
                                        countryterritoryCode             AFG
                                                 popData2019        38041757
                                                continentExp            Asia
  Cumulative_number_for_14_days_of_COVID-19_cases_per_100000      1.69813397
```

<!-- grade_id: cell-abde9c76b50503ba, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert isinstance(LoD, list)
assert len(LoD) == 44136
assert isinstance(LoD[0], dict)
```

**Q7:** With the successfully created `LoD`, create a pandas DataFrame named `covid_df`.

<!-- grade_id: cell-5645b511e37e6fd0, points: 1, grade: true, solution: true -->
```python
covid_df = pd.DataFrame(LoD)

# Display a subset of the data frame
covid_df.head()
```

Output:
```
dateRep day month  year  cases  deaths countriesAndTerritories geoId  \
0  21/09/2020  21    09  2020      0       0             Afghanistan    AF   
1  20/09/2020  20    09  2020    125       4             Afghanistan    AF   
2  19/09/2020  19    09  2020     47       1             Afghanistan    AF   
3  18/09/2020  18    09  2020      0       0             Afghanistan    AF   
4  17/09/2020  17    09  2020     17       0             Afghanistan    AF   

  countryterritoryCode  popData2019 continentExp  \
0                  AFG   38041757.0         Asia   
1                  AFG   38041757.0         Asia   
2                  AFG   38041757.0         Asia   
3                  AFG   38041757.0         Asia   
4                  AFG   38041757.0         Asia   

  Cumulative_number_for_14_days_of_COVID-19_cases_per_100000  
0                                         1.69813397          
1                                         1.89265706          
2                                  
```

<!-- grade_id: cell-abde9c76b50503ba, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert covid_df.shape == (44136, 12)
```

> You've reached the second checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 2: The data frame `covid_df` is not tidy.  What are two reasons it's not tidy, and how would you fix these issues?  (Note: you don't have to actually fix them, just make a plan.)

dataRep and day, month, year represent the same thing 
countriesAndTerritories geoId countryterritoryCode

two mappings in the same table

---

**Q8:** Find the number of unique countries, assigning to `ncountry`

<!-- grade_id: cell-e292daf935471ed3, points: 1, grade: true, solution: true -->
```python
ncountry = covid_df["countriesAndTerritories"].nunique()

# Display the result
print(ncountry)
```

Output:
```
210
```

**Q9:** Use a GroupBy to group by `"countryterritoryCode"` and then aggregate, computing the `max` and `mean` for the columns `cases` and `deaths`.  Name the resultant dataframe `covid_df_agg`.

<!-- grade_id: cell-aa5f6ab036596702, points: 2, grade: true, solution: true -->
```python
covid_df1 = covid_df.groupby("countryterritoryCode")
aggs = {"cases":["max", "mean"],"deaths":["max", "mean"]}
covid_df_agg = covid_df1.agg(aggs)
# Display the aggregated data frame
covid_df_agg

```

Output:
```
cases              deaths           
                        max         mean    max       mean
countryterritoryCode                                      
ABW                     176    19.404372      2   0.125683
AFG                    1063   152.515625     56   5.628906
AGO                     180    21.690217      8   0.826087
AIA                       2     0.016760      0   0.000000
ALB                     178    62.868020      6   1.837563
...                     ...          ...    ...        ...
XKX                     496    78.453125     40   3.192708
YEM                     116    12.303030     52   3.551515
ZAF                   13944  3339.449495    572  80.570707
ZMB                     527    75.566845     40   1.764706
ZWE                     490    41.529730     18   1.216216

[209 rows x 4 columns]
```

> You've reached the third (and final) checkpoint in the lab.  Make sure to have it signed off by the instructor or TA.
>
> Checkpoint 3: Use `covid_df_agg` to answer the following questions:
>
> - For the timeframe represented by this dataset, which country had the highest maximum death count?
> - For the timeframe represented by this dataset, of the countries with a maximum of at least 10 deaths, which country had the lowest mean case count?

```python
hi = covid_df_agg.loc[covid_df_agg["deaths"]["max"]== covid_df_agg["deaths"]["max"].max(),:]
print(hi)
hii = covid_df_agg.loc[covid_df_agg["deaths"]["max"]>=10]

hiii = hii.loc[hii["cases"]["mean"]== hii["cases"]["mean"].min()]
hiii

```

Output:
```
cases               deaths            
                        max          mean    max        mean
countryterritoryCode                                        
USA                   78427  25582.007519   4928  750.033835
                     cases           deaths          
                       max      mean    max      mean
countryterritoryCode                                 
TCD                     83  6.188172     10  0.435484
```

---

---
## Part C

How much time (in minutes/hours) did you spend on this lab outside of class?

45 mins 