# Denison DA210/CS181 Homework 5.b - Step 1

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

## Part A

**Q1:** One of the reasons for the existence of *Unicode* is its ability to use strings that go beyond the limitations of the keyboard.  Relative to the discussion in the chapter, *Unicode* is about the *strings* we can use in our programs, and the issue of how they translate/map to a sequence of *bytes* (i.e. their *encoding*) is a separate concept.

When we have the *code-point* (generally a hex digit sequence identifying an index into the set of characters) for a Unicode character that is beyond our normal keyboard characters, we can include them in our strings by using the `\u` escape prefix followed by the hex digits for the code-point.  Consider the Python string s:

    s = "Unicode examples: \u2B2C and \u266A and \u1F60 and " \
        "\u265E and \u0394 and \u0402"
        
Write code to print `s`, then assign to `b8` the UTF-8 encoding of `s`, and `b16` the UTF-16BE encoding of `s`.  For each, use the `.hex()` method of the `bytes` data type to see a hex version of the encoded values.

<!-- grade_id: cell-e2e0441cf1a37b54, points: 2, grade: true, solution: true -->
```python
s = "Unicode examples: \u2B2C and \u266A and \u1F60 and " \
    "\u265E and \u0394 and \u0402"
    
b8 = s.encode("utf-8")

b16 = s.encode("utf-16")

print(len(b8.hex()))
print(len(b16.hex()))
print(len(b8))
print(len(b16))
print(len(s))

```

Output:
```
118
200
59
100
49
```

**Q2:** Consider the results of Q1, and answer the following questions: 

- which of the hex representations is longer?
- give explicit lengths for `b8`, `b16`, and for the two `hex()` transformations.
- how does this compare to the length of `s`?

<!-- grade_id: cell-607a3db88789bed0, points: 1, grade: true, solution: true -->
The hex representation b16 ("utf-16") of s is longer. The lengths of b8, b16, b8.hex(), b16.hex() are 118, 200, 59, 100, respectively. All of these representations of s are longer than the length of s itself. 

**Q3:** Write a function
```
    shiftLetter(letter, n)
```
whose parameter, `letter` should be a single character.  If the character is between `"A"` and `"Z"`, the function returns an upper case character $n$ positions further along, and "wrapping" if the + $n$ mapping goes past `"Z"`.  Likewise, it should map the lowercase characters between `"a"` and `"z"`.  If the parameter `letter` is anything else, or not of length 1, the function should return `letter`.

Hint: review functions `ord()` and `chr()` from Section 21.1.1, as well as the modulus operator `%`.

<!-- grade_id: cell-fb6f443b7ea95ec8, points: 0, grade: false, solution: true -->
```python
def shiftLetter(letter, n): 
    if len(letter) != 1:
        return letter
    num = ord(letter)
    if num in range(65, 91) or num in range(97, 123):
        if num  == 90 or num == 122:
            num = num - 26
        dec = num + n 
        if dec >= 91 and dec < 97: 
            dec = dec - 26
        elif dec > 97 and dec > 122:
            dec = dec - 26
        new = chr(dec) 
    else: 
        return letter
    return new
```

```python
# Debugging cell
print(shiftLetter("A", 1))
print(shiftLetter("Z", 1))
print(shiftLetter("M", 3))
print(shiftLetter("X", 2))
```

Output:
```
B
A
P
Z
```

<!-- grade_id: cell-8d8fb21572039cbd, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert shiftLetter("A", 1) == "B"
assert shiftLetter("Z", 1) == "A"
assert shiftLetter("M", 3) == "P"
assert shiftLetter("X", 2) == "Z"

assert shiftLetter("a", 1) == "b"
assert shiftLetter("z", 1) == "a"
assert shiftLetter("m", 3) == "p"
assert shiftLetter("x", 2) == "z"


assert shiftLetter(" ", 1) == " "
assert shiftLetter("-", 1) == "-"
assert shiftLetter("/", 3) == "/"
assert shiftLetter("^", 5) == "^"

assert shiftLetter("abc", 4) == "abc"
```

**Q4:** Building on the previous exercise, write a function
```
    encrypt(plaintext, n)
```
that performs a `shiftLetter` for each of the letters in `plaintext` and accumulates and returns the resultant string.

<!-- grade_id: cell-32d65f646b2da222, points: 0, grade: false, solution: true -->
```python
def encrypt(plaintext, n):
    hi = ""
    for letter in plaintext:
        env = shiftLetter(letter, n)
        hi = hi + env
    return hi
```

```python
# Debugging cell
print(encrypt("abc", 1))
print(encrypt("xyz XYZ", 2))
```

Output:
```
bcd
zab ZAB
```

<!-- grade_id: cell-fef1d4cb2dc5b399, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert encrypt("abc", 1) == "bcd"
assert encrypt("xyz XYZ", 2) == "zab ZAB"
```

**Q5:** Write a function 
```
    singleByteChars(s)
   ``` 
that takes its argument, `s`, and returns `True` if all the characters in `s` can be encoded by a single byte, and `False` otherwise.

<!-- grade_id: cell-938f90290ee40c69, points: 0, grade: false, solution: true -->
```python
def singleByteChars(s): 
    for l in s:
        c = hex(ord(l))
        if len(c) > 4: 
            return False 
    return True


print(singleByteChars("Hello \u2B2C"))
singleByteChars("Hello")
```

Output:
```
False
True
```

<!-- grade_id: cell-a02544e577b74f0a, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert singleByteChars("Hello \u2B2C") == False
assert singleByteChars("Hello") == True
```

**Q6:** Suppose you have, in your Python program, a variable that refers to a `bytes` data type, like `mystery` refers to the `bytes` constant literal as given here:
```
    mystery = b'\xc9\xa2\x95}\xa3@\x89\xa3@\x87\x99\x85'    \
              b'\x81\xa3@\xa3\x96@\x82\x85@\xa2\x96\x93'    \
              b'\xa5\x89\x95\x87@\x97\x99\x96\x82\x93\x85'  \
              b'\x94\xa2o@@\xe8\x96\xa4@\x82\x85\xa3Z'
```              
Perhaps this value came from a network message, or from a file.  But you suspect that it, in fact, holds the bytes for a character string, and you need to figure out how it is encoding.  Assume that you have narrowed the encodings down to one of the following:

- 'UTF-8'
- 'UTF-16BE'
- 'cp037'
- 'latin_1'

Write code to convert the byte sequence to a character string, and determine the correct encoding.  By the end of your code, assign to `s_mystery` the correctly decoded string.

<!-- grade_id: cell-6ce4980e23e78176, points: 0, grade: false, solution: true -->
```python
mystery = b'\xc9\xa2\x95}\xa3@\x89\xa3@\x87\x99\x85'    \
          b'\x81\xa3@\xa3\x96@\x82\x85@\xa2\x96\x93'    \
          b'\xa5\x89\x95\x87@\x97\x99\x96\x82\x93\x85'  \
          b'\x94\xa2o@@\xe8\x96\xa4@\x82\x85\xa3Z'
          
s1 = mystery.decode("utf-16be")


s_mystery = mystery.decode("cp037")

s3 = mystery.decode("latin_1")


```

---

---

## Part B

**Q7:** How much time (in minutes/hours) did you spend on this homework assignment?

50 minutes

**Q8:** Who was your partner for this assignment?  If you worked alone, say so instead.

Alone